### AUTOEVALUACIÓN LISTAS Y TUPLAS

**Ejercicio 1**

Defina una tupla llamada _mi\_lista_ y una tupla llamada _mi\_tupla_ que contengan ambas los siguientes valores:
- En primer lugar el valor 10.
- En segunda lugar la cadena 'segundo elemento' (cuidado con los espacios).
- El booleano False.

```python
# COMPLETA EL EJERCICIO 1 AQUI
# Descomenta las siguientes líneas y defina mi_lista y mi_tupla con el siguiente contenido:
    # - 10 (número entero)
    # - "segundo elemento" (cadena)
    # - False (Boolean)

# mi_lista = []
# mi_tupla = ()
```

<!-- grade_id: cell-fee17c887a202357, points: 0, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
### BEGIN HIDDEN TESTS
assert(mi_lista == [10, 'segundo elemento', False])
assert(mi_lista[0] == 10)
assert(mi_lista[1] == 'segundo elemento')
assert(mi_lista[2] == False)
assert(len(mi_lista) == 3)

assert(mi_tupla == (10, 'segundo elemento', False))
assert(mi_tupla[0] == 10)
assert(mi_tupla[1] == 'segundo elemento')
assert(mi_tupla[2] == False)
assert(len(mi_tupla) == 3)
### END HIDDEN TESTS
```

**Ejercicio 2**

Completa la siguiente función la cual devuelva el último elemento de una lista pasada por parametro. Cíñete a los contenidos dados hasta ahora, no es necesario hacer comprobaciones de tamaño de la lista. Utiliza el acceso mediante índices para resolverlo:

<!-- grade_id: cell-880d2fb8f63ab428, points: 0, grade: false, solution: true -->
```python
# COMPLETA EL EJERCICIO 2 AQUI

def ultimo_elemento(lista):
    ### BEGIN SOLUTION
    return lista[-1]
    ### END SOLUTION
```

<!-- grade_id: cell-5b79c4b39c5b7e39, points: 5, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
### BEGIN HIDDEN TESTS
test_list1 = ['hello']
test_list2 = []

assert(ultimo_elemento(test_list1) == 'hello')

for i in range(10):
    test_list1.append(i)
    assert(ultimo_elemento(test_list1) == i)

test_list2.extend(test_list1)
assert(ultimo_elemento(test_list1) == 9)

test_list2.append('goodbye')
assert(ultimo_elemento(test_list2) == 'goodbye')
### END HIDDEN TESTS
```

**Ejercicio 3**

Completa la siguiente función la cual devuelva el segundo, tercer y cuarto elemento de una lista pasada por parámetro. Cíñete a los contenidos dados hasta ahora, no es necesario hacer comprobaciones de tamaño de la lista. Utiliza recortes para resolverlo:

<!-- grade_id: cell-52e5ffe77fc608f1, points: 0, grade: false, solution: true -->
```python
# COMPLETA EL EJERCICIO 3 AQUI

def recorte_lista(lista):
    ### BEGIN SOLUTION
    return lista[1:4]
    ### END SOLUTION
```

<!-- grade_id: cell-1c41e123d6e12696, points: 5, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
### BEGIN HIDDEN TESTS
test_list1 = []
assert(recorte_lista(test_list1) == [])

for i in range(5):
    test_list1.append(i)
    if i == 0:  
        assert(recorte_lista(test_list1) == [])
    elif i == 1:
        assert(recorte_lista(test_list1) == [i])
    elif i == 2:
        assert(recorte_lista(test_list1) == [i-1,i])
    else:
        assert(recorte_lista(test_list1) == [1,2,3])

test_list1.reverse()
assert(recorte_lista(test_list1) == [3,2,1])
### END HIDDEN TESTS
```

**Ejercicio 4**

Mediante el uso de métodos de listas, completa la siguiente función que haga las siguientes operaciones en orden:
- Añada al final de lista1 el contenido de lista2
- Añada el elemento _element_ al final de lista1
- Extraiga el elemento en la posición 3

```python
# COMPLETA EL EJERCICIO 4 AQUI

def prueba_metodos(lista1,lista2,element):
    # en la siguiente línea añade a lista1 el contenido de lista2
    
    # en la siguiente línea añade element al final de lista1
    
    # en la siguiente línea extrae el elemento en la posición 3
    
    # no edite este return
    return lista1


```

<!-- grade_id: cell-3a804c0c8b1ba3d4, points: 5, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
### BEGIN HIDDEN TESTS
assert(prueba_metodos([1,2],[],'element') == [1, 2])
assert(prueba_metodos([],[1,2],'element') == [1, 2])
assert(prueba_metodos([1],[1],'element') == [1,1])
assert(prueba_metodos([1,2],[1],'element') == [1, 2, 'element'])
assert(prueba_metodos([1],[1,2],'element') == [1, 1, 'element'])
assert(prueba_metodos([1,2],[1,2],'element') == [1, 2, 2, 'element'])
assert(prueba_metodos([1,2,3],[1,2],'element') == [1, 2, 1, 2, 'element'])
assert(prueba_metodos([1,2],[1,2,3],'element') == [1, 2, 2, 3, 'element'])
### END HIDDEN TESTS
```

```python

```