### AUTOEVALUACIÓN INTRODUCCIÓN A PYTHON

**Ejercicio 1**

Realiza la siguiente asignación de variables:
- Una variable numérica entera de nombre _mi\_entero_ con el valor 100.

- Una variable numérica en octal de nombre _mi\_octal_ con valor 20 (16 en decimal).

- Una cadena de nombre _mi\_cadena_ con el contenido "Hola Python", asegurate de respetar mayusculas y el espacio intermedio y de no dejar espacios al principio y al final.

- Una variable booleana de nombre _mi\_bool_ con valor True.

Respeta los nombres de las variables y los valores pedidos:

```python
# COMPLETA EL EJERCICIO 1 AQUI
# Descomenta la siguiente linea y completa tu variable entera:
mi_entero = 
# Descomenta la siguiente linea y completa tu variable octal:
mi_octal = 
# Descomenta la siguiente linea y completa tu cadena:
mi_cadena = 
# Descomenta la siguiente linea y completa tu booleano:
mi_bool = 


```

Output:
```
Input In [1]
    mi_entero =
                ^
SyntaxError: invalid syntax
```

```python
# Escribe debajo de este comentario tu variable entera:
mi_entero = 100
# Escribe debajo de este comentario tu variable octal:
mi_octal = 0o20
# Escribe debajo de este comentario tu cadena:
mi_cadena = "Hola Python"
# Escribe debajo de este comentario tu booleano:
mi_bool = True
```

<!-- grade_id: cell-bf76ba2e7507e82d, points: 5, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
# BEGIN HIDDEN TESTS 
assert(mi_entero == 100)
assert(mi_octal  == 16)
assert(mi_octal  == 0o20)
assert(mi_cadena == "Hola Python")
assert(mi_bool)
# END HIDDEN TESTS 
```

**Ejercicio 2**

Dado el radio de la base de un cilindro y su altura (en cm) calcula el area de su base y posteriormente el volumen de dicho cilindro:

```python
# COMPLETA EL EJERCICIO 2 AQUI
radio = 6
pi = 3.14
altura = 20

# Descomenta la siguiente linea y calcula el area de la base
# area_base = 

# Descomenta la siguiente linea y calcula el volumen usando el area de la base calculado anteriormente
# volumen = 
```

```python
# COMPLETA EL EJERCICIO 2 AQUI
radio = 6
pi = 3.14
altura = 20

# Descomenta la siguiente linea y calcula el area de la base
area_base = pi * (radio**2)

# Descomenta la siguiente linea y calcula el volumen usando el area de la base calculado anteriormente
volumen = area_base * altura
```

<!-- grade_id: cell-0afd55718551d970, points: 5, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
# BEGIN HIDDEN TESTS 
assert(area_base == 113.04)
assert(volumen  == 2260.8)
# END HIDDEN TESTS 
```

**Ejercicio 3**

Dada una variable _A_ con un valor asignado, y una variable _B_ con otro valor asignado, intercambia el valor de estas dos variables de manera que el contenido de _A_ pase a  _B_  y viceversa mediante el uso de otra variable externa llamada _C_:

```python
# COMPLETA EL EJERCICIO 3 AQUI
A = 10
B = 50

# Realiza el intercambio a continuación usando una variable de nombre C (descomentala):
# C = 
```

```python
# COMPLETA EL EJERCICIO 3 AQUI
A = 10
B = 50

# Realiza el intercambio a continuación usando una variable de nombre C:
C = B
B = A
A = C
```

<!-- grade_id: cell-c2e4ef6d83ac395c, points: 5, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
assert(A == 50)
assert(B == 10)
assert(C == 50 or C == 10)
```

**Ejercicio 4**

Coloca paréntesis () en la siguiente expresión para que se corresponda con la fórmula del area de un trapecio. Recuerda que el area del trapecio es la suma de las bases por la altura dividido entre dos:

```python
# COMPLETA EL EJERCICIO 3 AQUI
base_mayor = 10
base_menor = 4
altura_tr = 4
# Coloca los parentesis de forma correcta:
area_tr = base_mayor + base_menor * altura_tr  / 2
```

```python
# COMPLETA EL EJERCICIO 3 AQUI
base_mayor = 10
base_menor = 4
altura_tr = 4
# Coloca los parentesis de forma correcta:
area_tr = ((base_mayor + base_menor) * altura_tr ) / 2
```

<!-- grade_id: cell-5b5576449acf44b4, points: 5, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
# BEGIN HIDDEN TESTS 
assert(area_tr == 28.0)
# END HIDDEN TESTS 
```

```python

```