### AUTOEVALUACIÓN CONTROL DE FLUJO


**Ejercicio 1**

Completa la siguiente función la cual comprueba si una persona es mayor de edad o no. En caso de ser mayor de edad devolverá True y sino False. La funcion NO ha de calcular si la edad introducida es mayor que 0:

<!-- grade_id: cell-91b62f528ca98b71, points: 0, grade: false, solution: true -->
```python
# COMPLETA EL EJERCICIO 1 AQUI

def es_mayor_edad(edad):
    # YOUR CODE HERE
    if edad < 18:
        return False
    else:
        return True
```

<!-- grade_id: cell-d607e8394b4add72, points: 5, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
### BEGIN HIDDEN TESTS
assert(not es_mayor_edad(0))
assert(not es_mayor_edad(-10))
assert(not es_mayor_edad(3.2))
assert(not es_mayor_edad(17))
assert(es_mayor_edad(18))
assert(es_mayor_edad(50.8))
assert(es_mayor_edad(0o200))
assert(not es_mayor_edad(True))
assert(not es_mayor_edad(True))
### END HIDDEN TESTS
```

**Ejercicio 2**

Completa la siguiente función la cual define la nota que ha sacado un estudiante según la calificación obtenida.
- Si su calificación está entra 0 y 5 (este útimo no incluido) devolverá la cadena 'Suspenso'.
- Si su calificación está entra 5 y 6 (este útimo no incluido) devolverá la cadena 'Suficiente'.
- Si su calificación está entra 6 y 7 (este útimo no incluido) devolverá la cadena 'Bien'.
- Si su calificación está entra 7 y 9 (este útimo no incluido) devolverá la cadena 'Notable'.
- Si su calificación está entra 9 y 10 (este útimo incluido) devolverá la cadena 'Sobresaliente'.
- Si la calificación es menor que 0 o mayor que 10 devolverá la cadena 'ERROR'.

<!-- grade_id: cell-e390abb46479f273, points: 0, grade: false, solution: true -->
```python
# COMPLETA EL EJERCICIO 2 AQUI

def define_nota(calif):
    if calif >= 0 and calif < 5:
        return "Suspenso"
    elif calif >= 5 and calif < 6:
        return "Suficiente"
    elif calif >= 6 and calif < 7:
        return "Bien"
    elif calif >= 7 and calif < 9:
        return "Notable"
    elif calif >= 9 and calif <= 10:
        return "Sobresaliente"
    else:
        return "ERROR"
```

<!-- grade_id: cell-f46c4b2811513ef3, points: 5, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
### BEGIN HIDDEN TESTS
assert(define_nota(-1) == 'ERROR')
assert(define_nota(-10) == 'ERROR')
assert(define_nota(0) == 'Suspenso')
assert(define_nota(4.9) == 'Suspenso')
assert(define_nota(5) == 'Suficiente')
assert(define_nota(5.9) == 'Suficiente')
assert(define_nota(6) == 'Bien')
assert(define_nota(6.9) == 'Bien')
assert(define_nota(7) == 'Notable')
assert(define_nota(8.9) == 'Notable')
assert(define_nota(9) == 'Sobresaliente')
assert(define_nota(10) == 'Sobresaliente')
assert(define_nota(10.1) == 'ERROR')
assert(define_nota(100) == 'ERROR')
### END HIDDEN TESTS
```

**Ejercicio 3**

Completa la siguiente función la cual recibe un número y devuelve una lista con todos los números impares desde el 1 hasta el enviado ambos incluidos. Si el número es menor o igual que cero y mayor que 20 devolverá una lista vacía []. Usa un bucle while para resolverlo y recuerda los métodos dados en el tema de listas para ayudarte a resolverlo:

<!-- grade_id: cell-3e195ae40143a93e, points: 0, grade: false, solution: true -->
```python
# COMPLETA EL EJERCICIO 3 AQUI

def impares(num):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6e6fcd2de49ab52f, points: 5, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
### BEGIN HIDDEN TESTS
assert(impares(0) == [])
assert(impares(-1) == [])
assert(impares(-5) == [])
assert(impares(21) == [])
assert(impares(1) == [1])
assert(impares(2) == [1])
assert(impares(5) == [1,3,5])
assert(impares(20) == [1,3,5,7,9,11,13,15,17,19])
### END HIDDEN TESTS
```

Output:
```
Input In [6], in <cell line: 3>()
      1 # NO EDITE ESTA CELDA
      2 ### BEGIN HIDDEN TESTS
----> 3 assert(impares(0) == [])
      4 assert(impares(-1) == [])
      5 assert(impares(-5) == [])

Input In [5], in impares(num)
      3 def impares(num):
      4     # YOUR CODE HERE
----> 5     raise NotImplementedError()

NotImplementedError:
```

**Ejercicio 4**

Completa la siguiente función la cual recibe una lista de números enteros y calcula la suma (sumatorio) de todos ellos en valor absoluto. Si la lista está vacia devolverá 0. Usa un bucle for para resolverlo:

<!-- grade_id: cell-73426cdd01a742d7, points: 0, grade: false, solution: true -->
```python
# COMPLETA EL EJERCICIO 4 AQUI

def sumatorio_absoluto(lista):
    # YOUR CODE HERE
    sumatorio = 0
    for i in lista:
        if i < 0:
            sumatorio = sumatorio - i
        else:
            sumatorio = sumatorio + i
    
    return sumatorio
```

<!-- grade_id: cell-fd51adb56b2ae36c, points: 5, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
### BEGIN HIDDEN TESTS
assert(sumatorio_absoluto([]) == 0)
assert(sumatorio_absoluto([1]) == 1)
assert(sumatorio_absoluto([-1]) == 1)
assert(sumatorio_absoluto([1,4,6]) == 11)
assert(sumatorio_absoluto([-1,4,-6]) == 11)
### END HIDDEN TESTS
```

```python

```