### AUTOEVALUACIÓN FUNCIONES

**Ejercicio 1**

Define una función llamada _area\_circulo_ la cual calcule el area de la base de un cilindro (círculo) dado el rádio por parámetro y, en caso de no ser pasado este parametro, que telga el valor 10 por defecto:

<!-- grade_id: cell-3a98f8ea28d51d44, points: 0, grade: false, solution: true -->
```python
# COMPLETA EL EJERCICIO 1 AQUI
# Define la función a continuación, usa el valor de pi como 3.14:

def area_circulo(radio):
    # YOUR CODE HERE
    return 3.14 *(radio**2)
```

<!-- grade_id: cell-954952c2cd354a46, points: 5, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
# BEGIN HIDDEN TESTS 
assert(area_circulo(0) == 0.0)
assert(area_circulo(1) == 3.14)
assert(area_circulo(10) == 314.0)
assert(area_circulo(100) == 31400.0)
assert(area_circulo(-10) == 314.0)
assert(area_circulo(-5) == 78.5)
assert(area_circulo(True) == 3.14)
assert(area_circulo(False) == 0.0)
# END HIDDEN TESTS 
```

**Ejercicio 2**

Define una función llamada _volumen\_cilindro_ la cual calcule el volumen de un cilindro dado el rádio de la base y su altura por parámetro en ese orden. Utiliza la función _area\_circulo_ definida en el anterior ejercicio para calcular el área de la base dentro de la función volumen:

<!-- grade_id: cell-2e3314a33e65af00, points: 0, grade: false, solution: true -->
```python
# COMPLETA EL EJERCICIO 2 AQUI

def volumen_cilindro(radio, altura):
    # YOUR CODE HERE
    return area_circulo(radio) * altura
```

<!-- grade_id: cell-3c1157e43302672a, points: 5, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
# BEGIN HIDDEN TESTS 
assert(volumen_cilindro(0, 0) == 0.0)
assert(volumen_cilindro(1, 0) == 0.0)
assert(volumen_cilindro(0, 1) == 0.0)
assert(volumen_cilindro(10, 1) == 314.0)
assert(volumen_cilindro(10, 10) == 3140.0)
assert(volumen_cilindro(-1, 10) == 31.400000000000002)
assert(volumen_cilindro(10, -1) == -314.0)
assert(volumen_cilindro(True, True) == 3.14)
assert(volumen_cilindro(True, False) == 0.0)
assert(volumen_cilindro(False, True) == 0.0)
assert(volumen_cilindro(False, False) == 0.0)
# END HIDDEN TESTS 
```

**Ejercicio 3**

Llama a la función _volumen\_cilindro_ creada anteriormente para calcular el volumen de un cilindro dados un rádio y una altura:

```python
# COMPLETA EL EJERCICIO 3 AQUI
radio_base = 5
altura_cilindro = 10

# descomenta la siguiente línea llamando a tu función para asignar a la variable solución el volumen del cilindro:
# solucion = 
```

<!-- grade_id: cell-d57386c63deb464d, points: 5, grade: true, solution: false -->
```python
# CELDA DE TESTS TEMPORAL
# BEGIN HIDDEN TESTS
assert(solucion == 785.0)
# END HIDDEN TESTS
```

Output:
```
NameError                                 Traceback (most recent call last)
Input In [6], in <cell line: 3>()
      1 # CELDA DE TESTS TEMPORAL
      2 # BEGIN HIDDEN TESTS
----> 3 assert(solucion == 785.0)

NameError: name 'solucion' is not defined
```

```python

```