### AUTOEVALUACIÓN ORIENTACIÓN A OBJETOS

**Ejercicio 1**

Completa la clase Aula la cual representa un aula de un colegio. Esta tendrá dos atributos de clase: los alumnos que la forman y el estado de la clase el cual representará si la clase está siendo impartida o no. El valor de estado al instanciarse la clase ha de ser False. Completa su constructor, los metodos getter y setter para el atributo alumno y los métodos para la gestión del estado de la clase: 

```python
# COMPLETA EL EJERCICIO 1 AQUI

class Aula:
    """Aula del colegio."""
    
    def __init__(self, alumnos):
        # Complete el método
        
    def getAlumnos(self):
        # Complete el método
                
    def setAlumnos(self, alumnos):
        # Complete el método
    
    def empiezaClase(self):
        # Complete el método
        
    def acabaClase(self):
        # Complete el método
        
    def getEstado(self):
       # Complete el método
    
    

```

<!-- grade_id: cell-e2933563594af41d, points: 10, grade: true, solution: false -->
```python
# NO EDITE ESTA CELDA
### BEGIN HIDDEN TESTS
mi_aula = Aula(20)
assert(mi_aula.getEstado() == False)
assert(mi_aula.getAlumnos() == 20)
mi_aula.setAlumnos(24)
assert(mi_aula.getAlumnos() == 24)
mi_aula.empiezaClase()
assert(mi_aula.getEstado() == True)
mi_aula.acabaClase()
assert(mi_aula.getEstado() == False)
### END HIDDEN TESTS
```

```python

```